# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    a = list(x)
    b = list(y)
    
    if len(a) != len(b):
        return False
        
    for i in range(len(a)):
        if a[i] not in b:
            return False
        
    return True
    
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
def find_small():
    s = True
    x = 1

    while s:
        if same_digits(str(2*x), str(3*x)) and same_digits(str(3*x), str(4*x)):
            if same_digits(str(4*x), str(5*x)) and same_digits(str(5*x),str(6*x)):
                return x
        else:
            x += 1

    return 0

```

```python
find_small()
```

Output:
```
142857
```

```python

```