# Character Entropy

In this notebook, you will learn about information entropy and how to compute it for discrete probability distributions.

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import seaborn as sns
import numpy as np
```

```python
from ipywidgets import interact
```

## Character counting and entropy

Write a function, `char_probs`, that takes a string and computes the probabilities of each character in the string:

* First do a character count and store the result in a dictionary.
* Then divide each character count by the total number of character to compute the normalized probabilties.
* Return the dictionary of characters (keys) and probabilities (values).

<!-- grade_id: charentropy-a, points: 0, grade: false, solution: true -->
```python
def char_probs(s):
    """Find the probabilities of the unique characters in the string s.
    
    Parameters
    ----------
    s : str
        A string of characters.
    
    Returns
    -------
    probs : dict
        A dictionary whose keys are the unique characters in s and whose values
        are the probabilities of those characters.
    """
    dic = {}
    count = len(s)
    
    for ele in s:
        if ele in dic:
            dic[ele] += 1
        else:
            dic[ele] = 1
    
    for keys in dic:
        dic[keys] /= count
        
    return dic
        
```

<!-- grade_id: charentropy-b, points: 2, grade: true, solution: false -->
```python
test1 = char_probs('aaaa')
assert np.allclose(test1['a'], 1.0)
test2 = char_probs('aabb')
assert np.allclose(test2['a'], 0.5)
assert np.allclose(test2['b'], 0.5)
test3 = char_probs('abcd')
assert np.allclose(test3['a'], 0.25)
assert np.allclose(test3['b'], 0.25)
assert np.allclose(test3['c'], 0.25)
assert np.allclose(test3['d'], 0.25)
```

The [entropy](http://en.wikipedia.org/wiki/Entropy_%28information_theory%29) is a quantiative measure of the disorder of a probability distribution. It is used extensively in Physics, Statistics, Machine Learning, Computer Science and Information Science. Given a set of probabilities $P_{i}$, the entropy is defined as:

$$H = - \Sigma_i P_i \log_2(P_i)$$ 

In this expression $\log_2$ is the base 2 log (`np.log2`), which is commonly used in information science. In Physics the natural log is often used in the definition of entropy.

Write a funtion, `entropy`, that computes the entropy of a probability distribution. The probability distribution will be passed as a Python `dict`. The keys in the `dict` will be the unique values in the distribution and the values in the `dict` will be the probabilities for those keys.

To compute the entropy, you should:

* First convert the values (probabilities) of the `dict` to a Numpy array of probabilities.
* Then use other Numpy functions (`np.log2`, etc.) to compute the entropy.
* Don't use any `for` or `while` loops in your code.

<!-- grade_id: charentropy-c, points: 0, grade: false, solution: true -->
```python
def entropy(d):
    """Compute the entropy of a dict d whose values are probabilities."""
    arr = [d[ele] for ele in d]
    
    r = np.log2(arr)* arr*-1
    r = r.sum()
    return r
```

```python
entropy({'a': 0.5, 'b': 0.5})
```

Output:
```
1.0
```

<!-- grade_id: charentropy-d, points: 2, grade: true, solution: false -->
```python
assert np.allclose(entropy({'a': 0.5, 'b': 0.5}), 1.0)
assert np.allclose(entropy({'a': 1.0}), 0.0)
```

Use IPython's `interact` function to create a user interface that allows you to type a string into a text box and see the entropy of the character probabilities of the string. Try to get an intuitive sense of the entropy by typing in different strings. What strings make the entropy large? Small?

<!-- grade_id: charentropy-e, points: 2, grade: true, solution: true -->
```python
interact(char_probs, s=('aaabbb', 'addcccddee'))
```

Output:
```
{'a': 0.1, 'c': 0.3, 'd': 0.4, 'e': 0.2}
```

Find the entropies of the all of the books in the directory `/data/gutenberg`:

```python
ls /data/gutenberg/*.txt
```

Output:
```
/data/gutenberg/11.txt    /data/gutenberg/17208.txt  /data/gutenberg/33511.txt
/data/gutenberg/1342.txt  /data/gutenberg/23.txt     /data/gutenberg/84.txt
/data/gutenberg/1400.txt  /data/gutenberg/2701.txt
/data/gutenberg/1661.txt  /data/gutenberg/29021.txt
```

* Store your entropies in a list
* Create an appropriately labeled visualization of the distribution of entropies
* Print the mean entropy

<!-- grade_id: charentropy-f, points: 4, grade: true, solution: true -->
```python
l =[]
with open('/data/gutenberg/*.txt') as f:
    for in
    
```